```python
# Initialize Otter
import otter
grader = otter.Notebook("lab11.ipynb")
```

# Lab 11: SQL

In this lab, we are going to practice viewing, sorting, grouping, and merging tables with SQL. We will explore two datasets:
1. A "minified" version of the [Internet Movie Database](https://www.imdb.com/interfaces/) (IMDb). This SQLite database (~10MB) is a tiny sample of the much larger database (more than a few GBs). As a result, disclaimer that we may get wildly different results than if we use the whole database!

1. The money donated during the 2016 election using the [Federal Election Commission (FEC)'s public records](https://www.fec.gov/data/). You will be connecting to a SQLite database containing the data. The data we will be working with in this lab is quite small (~16MB); however, it is a sample taken from a much larger database (more than a few GBs).


**Due Date: Tuesday, November 8, 11:59 PM PT.**

### Lab Walk-Through
In addition to the lab notebook, we have also released a prerecorded walk-through video of the lab. We encourage you to reference this video as you work through the lab. Run the cell below to display the video.

**Note:** the walkthrough video is from Spring 2022, where this lab was labeled Lab 10.

```python
from IPython.display import YouTubeVideo
YouTubeVideo("uQ3E4pejmD8", list = 'PLQCcNQgUcDfpdBnhS-lPq8LPas48tkMgp', listType = 'playlist')
```

### Collaboration Policy
Data science is a collaborative activity. While you may talk with others about this assignment, we ask that you **write your solutions individually**. If you discuss the assignment with others, please **include their names** in the cell below.

**Collaborators:** *list names here*

```python
# Run this cell and the next one to set up your notebook
import numpy as np
import pandas as pd
import plotly.express as px
import sqlalchemy
from pathlib import Path
from zipfile import ZipFile

with ZipFile('data.zip', 'r') as zipObj:
   # Extract all the contents of zip file in current directory
   zipObj.extractall()

#Comment out this line after you have run it
!pip install ipython-sql

```

```python
# Just run this cell
%load_ext sql
```

## SQL Query Syntax

Throughout this lab, you will become familiar with the following syntax for the `SELECT` query:

```
SELECT [DISTINCT] 
    {* | expr [[AS] c_alias] 
    {,expr [[AS] c_alias] ...}}
FROM tableref {, tableref}
[[INNER | LEFT ] JOIN table_name
    ON qualification_list]
[WHERE search_condition]
[GROUP BY colname {,colname...}]
[HAVING search condition]
[ORDER BY column_list]
[LIMIT number]
[OFFSET number of rows];
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 0 [Tutorial]: Writing SQL in Jupyter Notebooks

### 1. `%%sql` cell magic

In lecture, we used the `sql` extension to call **`%%sql` cell magic**, which enables us to connect to SQL databses and issue SQL commands within Jupyter Notebooks.

Run the below cell to connect to a mini IMDb database.

```python
%sql sqlite:///imdbmini.db
```

<br/>

Above, prefixing our single-line command with `%sql` means that the entire line will be treated as a SQL command (this is called "line magic"). In this class we will most often write multi-line SQL, meaning we need "cell magic", where the first line has `%%sql` (note the double `%` operator).

The database `imdbmini.db` includes several tables, one of which is `Title`. Running the below cell will return first 5 lines of that table. Note that `%%sql` is on its own line.

We've also included syntax for single-line comments, which are surrounded by `--`.

```python
%%sql
/*
 * This is a
 * multi-line comment.
 */
-- This is a single-line/inline comment. --
SELECT *
FROM Name
LIMIT 5;
```

<br/><br/>

### 2. The Pandas command `pd.read_sql`

As of 2022, the `%sql` magic for Jupyter Notebooks is still in development (check out its [GitHub](https://github.com/catherinedevlin/ipython-sql). It is still missing many features that would justify real-world use with Python. In particular, its returned tables are *not* Pandas dataframes (for example, the query result from the above cell is missing an index).

The rest of this section describes how data scientists use SQL and Python in practice, using the Pandas command `pd.read_sql` ([documentation](https://pandas.pydata.org/docs/reference/api/pandas.read_sql.html)). **You will see both `%sql` magic and `pd.read_sql` in this course**. 


The below cell connects to the same database using the SQLAlchemy Python library, which can connect to several different database management systems, including sqlite3, MySQL, PostgreSQL, and Oracle. The library also supports an advanced feature for generating queries called an [object relational mapper](https://docs.sqlalchemy.org/en/latest/orm/tutorial.html) or ORM, which we won't discuss in this course but is quite useful for application development.

```python
# important!!! run this cell
import sqlalchemy

# create a SQL Alchemy connection to the database
engine = sqlalchemy.create_engine("sqlite:///imdbmini.db")
connection = engine.connect()
```

With the SQLAlchemy object `connection`, we can then call `pd.read_sql` which takes in a `query` **string**. Note the `"""` to define our multi-line string, which allows us to have a query span multiple lines. The resulting `df` DataFrame stores the results of the same SQL query from the previous section.

```python
# just run this cell
query = """
SELECT *
FROM Title
LIMIT 5;
"""

df = pd.read_sql(query, engine)
df
```

<br/>

**Long error messages**: Given that the SQL query is now in the string, the errors become more unintelligible. Consider the below (incorrect) query, which has a semicolon in the wrong place.

```python
# uncomment the below code and check out the error

# query = """
# SELECT *
# FROM Title;
# LIMIT 5
# """
# pd.read_sql(query, engine)
```

<br/>

Now that's an unruly error message!

<br/><br/>

### 3. A suggested workflow for writing SQL in Jupyter Notebooks

Which approach is better, `%sql` magic or `pd.read_sql`?

The SQL database generally contains much more data than what you would analyze in detail. As a Python-fluent data scientist, you will often query SQL databases to perform initial exploratory data analysis, a subset of which you load into Python for further processing.

In practice, you would likely use a combination of the two approaches. First, you'd try out some SQL queries with `%sql` magic to get an interesting subset of data. Then, you'd copy over the query into a `pd.read_sql` command for visualization, modeling, and export with Pandas, sklearn, and other Python libraries.

For SQL assignments in this course, to minimize unruly error messages while maximizing Python compatibility, we suggest the following "sandboxed" workflow:
1. Create a `%%sql` magic cell **below** the answer cell. You can copy in the below code:

    ```
    %% sql
    -- This is a comment. Put your code here... --
    ```
<br/>

1.  Work on the SQL query in the `%%sql` cell; e.g., `SELECT ... ;`
1. Then, once you're satisfied with your SQL query, copy it into the multi-string query in the answer cell (the one that contains the `pd.read_sql` call).

You don't have to follow the above workflow to get full credit on assignments, but we suggest it to reduce debugging headaches. We've created the scratchwork `%%sql` cells for you in this assignment, but **do not** add cells between this `%%sql` cell and the Python cell right below it. It will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file.


<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: The IMDb (mini) Dataset

Let's explore a miniature version of the [IMDb Dataset](https://www.imdb.com/interfaces/). This is the same dataset that we will use for the upcoming homework.


Let's load in the database in two ways (using both Python and cell magic) so that we can flexibly explore the database.

```python
engine = sqlalchemy.create_engine("sqlite:///imdbmini.db")
connection = engine.connect()
```

```python
%sql sqlite:///imdbmini.db
```

<br/>

```python
%%sql
SELECT * FROM sqlite_master WHERE type='table';
```

From running the above cell, we see the database has 4 tables: `Name`, `Role`, `Rating`, and `Title`.

<details>
    <summary>[<b>Click to Expand</b>] See descriptions of each table's schema.</summary>
    
**`Name`** – Contains the following information for names of people.
    
- nconst (text) - alphanumeric unique identifier of the name/person
- primaryName (text)– name by which the person is most often credited
- birthYear (integer) – in YYYY format
- deathYear (integer) – in YYYY format
    
    
**`Role`** – Contains the principal cast/crew for titles.
    
- tconst (text) - alphanumeric unique identifier of the title
- ordering (integer) – a number to uniquely identify rows for a given tconst
- nconst (text) - alphanumeric unique identifier of the name/person
- category (text) - the category of job that person was in
- characters (text) - the name of the character played if applicable, else '\\N'
    
**`Rating`** – Contains the IMDb rating and votes information for titles.
    
- tconst (text) - alphanumeric unique identifier of the title
- averageRating (real) – weighted average of all the individual user ratings
- numVotes (integer) - number of votes (i.e., ratings) the title has received
    
**`Title`** - Contains the following information for titles.
    
- tconst (text) - alphanumeric unique identifier of the title
- titleType (text) -  the type/format of the title
- primaryTitle (text) -  the more popular title / the title used by the filmmakers on promotional materials at the point of release
- isAdult (text) - 0: non-adult title; 1: adult title
- year (YYYY) – represents the release year of a title.
- runtimeMinutes (integer)  – primary runtime of the title, in minutes
    
</details>

<br/><br/>
From the above descriptions, we can conclude the following:
* `Name.nconst` and `Title.tconst` are primary keys of the `Name` and `Title` tables, respectively.
* that `Role.nconst` and `Role.tconst` are **foreign keys** that point to `Name.nconst` and `Title.tconst`, respectively.

<br/><br/>
<hr style="border: 1px solid #fdb515;" />

## Question 1

What are the different kinds of `titleType`s included in the `Title` table? Write a query to find out all the unique `titleType`s of films using the `DISTINCT` keyword.  (**You may not use `GROUP BY`.**)

```python
# %%sql
# /*
#  * Code in this scratchwork cell is __not graded.__
#  * Copy over any SQL queries you write here into the below Python cell.
#  * Do __not__ insert any new cells in between the SQL/Python cells!
#  * Doing so may break the autograder.
#  */
# -- Write below this comment. --

```

```python
query_q1 = """
...       # replace this with
...;      # your multi-line solution
"""

res_q1 = pd.read_sql(query_q1, engine)
res_q1
```

```python
grader.check("q1")
```

---

## Question 2

Before we proceed we want to get a better picture of the kinds of jobs that exist.  To do this examine the `Role` table by computing the number of records with each job `category`.  Present the results in descending order by the total counts.

The top of your table should look like this (however, you should have more rows):

| |category|total|
|-----|-----|-----|
|**0**|actor|21665|
|**1**|writer|13830|
|**2**|...|...|

```python
# %%sql
# /*
#  * Code in this scratchwork cell is __not graded.__
#  * Copy over any SQL queries you write here into the below Python cell.
#  * Do __not__ insert any new cells in between the SQL/Python cells!
#  * Doing so may break the autograder.
#  */
# -- Write below this comment. --

```

```python
query_q2 = """
...       # replace this with
...;      # your multi-line solution
"""

res_q2 = pd.read_sql(query_q2, engine)
res_q2
```

```python
grader.check("q2")
```

<br/>
If we computed the results correctly we should see a nice horizontal bar chart of the counts per category below:

```python
# just run this cell
px.bar(res_q2, x="total", y="category", orientation='h')
```

<br/><br/>
<hr style="border: 1px solid #fdb515;" />

## Question 3

Now that we have a better sense of the basics of our data, we can ask some more interesting questions.

The `Rating` table has the `numVotes` and the `averageRating` for each title. Which 10 films have the most ratings?

Write a SQL query that outputs three fields: the `title`, `numVotes`, and `averageRating` for the 10 films that have the highest number of ratings.  Sort the result in descending order by the number of votes.

*Hint*: The `numVotes` in the `Rating` table is not an integer! Use `CAST(Rating.numVotes AS int) AS numVotes` to convert the attribute to an integer.

```python
# %%sql
# /*
#  * Code in this scratchwork cell is __not graded.__
#  * Copy over any SQL queries you write here into the below Python cell.
#  * Do __not__ insert any new cells in between the SQL/Python cells!
#  * Doing so may break the autograder.
#  */
# -- Write below this comment. --

```

```python
query_q3 = """
...       # replace this with
...;      # your multi-line solution
"""


res_q3 = pd.read_sql(query_q3, engine)
res_q3
```

```python
grader.check("q3")
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Election Donations in New York City

Finally, let's analyze the Federal Election Commission (FEC)'s public records. We connect to the database in two ways (using both Python and cell magic) so that we can flexibly explore the database.

```python
# important!!! run this cell and the next one
import sqlalchemy
# create a SQL Alchemy connection to the database
engine = sqlalchemy.create_engine("sqlite:///fec_nyc.db")
connection = engine.connect()
```

```python
%sql sqlite:///fec_nyc.db
```

## Table Descriptions

Run the below cell to explore the **schemas** of all tables saved in the database.

If you'd like, you can consult the below linked FEC pages for the descriptions of the tables themselves.

* `cand` ([link](https://www.fec.gov/campaign-finance-data/candidate-summary-file-description/)): Candidates table. Contains names and party affiliation.
* `comm` ([link](https://www.fec.gov/campaign-finance-data/committee-summary-file-description/)): Committees table. Contains committee names and types.
* `indiv_sample_nyc` ([link](https://www.fec.gov/campaign-finance-data/contributions-individuals-file-description)): All individual contributions from New York City .

```python
%%sql
/* just run this cell */
SELECT sql FROM sqlite_master WHERE type='table';
```

<br/><br/>

Let's look at the `indiv_sample_nyc` table. The below cell displays individual donations made by residents of the state of New York. We use `LIMIT 5` to avoid loading and displaying a huge table.

```python
%%sql
/* just run this cell */
SELECT *
FROM indiv_sample_nyc
LIMIT 5;
```

You can write a SQL query to return the id and name of the first five candidates from the Democratic party, as below:

```python
%%sql
/* just run this cell */
SELECT cand_id, cand_name
FROM cand
WHERE cand_pty_affiliation = 'DEM'
LIMIT 5;
```

<br/><br/>
<hr style="border: 1px solid #fdb515;" />

## [Tutorial] Matching Text with `LIKE`

First, let's look at 2016 election contributions made by Donald Trump, who was a New York (NY) resident during that year. The following SQL query returns the `cmte_id`, `transaction_amt`, and `name` for every contribution made by any donor with "DONALD" and "TRUMP" in their name in the `indiv_sample_nyc` table.

Notes:
* We use the `WHERE ... LIKE '...'` to match fields with text patterns. The `%` wildcard represents at least zero characters. Compare this to what you know from regex!
* We use `pd.read_sql` syntax here because we will do some EDA on the result `res`.

```python
# just run this cell
example_query = """
SELECT 
    cmte_id,
    transaction_amt,
    name
FROM indiv_sample_nyc
WHERE name LIKE '%TRUMP%' AND name LIKE '%DONALD%';
"""

example_res = pd.read_sql(example_query, engine)
example_res
```

If we look at the list above, it appears that some donations were not by Donald Trump himself, but instead by an entity called "DONALD J TRUMP FOR PRESIDENT INC". Fortunately, we see that our query only seems to have picked up one such anomalous name.

```python
# just run this cell
example_res['name'].value_counts()
```

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

## Question 4

Revise the above query so that the 15 anomalous donations made by "DONALD J TRUMP FOR PRESIDENT INC" do not appear. Your resulting table should have 142 rows. 

Hints:
* Consider using the above query as a starting point, or checking out the SQL query skeleton at the top of this lab. 
* The `NOT` keyword may also be useful here.


```python
# %%sql
# /*
#  * Code in this scratchwork cell is __not graded.__
#  * Copy over any SQL queries you write here into the below Python cell.
#  * Do __not__ insert any new cells in between the SQL/Python cells!
#  * Doing so may break the autograder.
#  */
# -- Write below this comment. --

```

```python
query_q4 = """
...       # replace this with
...;      # your multi-line solution
"""

res_q4 = pd.read_sql(query_q4, engine)
res_q4
```

```python
grader.check("q4")
```

<br/><br/>
<hr style="border: 1px solid #fdb515;" />

## Question 5: `JOIN`ing Tables

Let's explore the other two tables in our database: `cand` and `comm`.

The `cand` table contains summary financial information about each candidate registered with the FEC or appearing on an official state ballot for House, Senate or President.

```python
%%sql
/* just run this cell */
SELECT *
FROM cand
LIMIT 5;
```

The `comm` table contains summary financial information about each committee registered with the FEC. Committees are organizations that spend money for political action or parties, or spend money for or against political candidates.

```python
%%sql
/* just run this cell */
SELECT *
FROM comm
LIMIT 5;
```

---

### Question 5a

Notice that both the `cand` and `comm` tables have a `cand_id` column. Let's try joining these two tables on this column to print out committee information for candidates.

List the first 5 candidate names (`cand_name`) in reverse lexicographic order by `cand_name`, along with their corresponding committee names. **Only select rows that have a matching `cand_id` in both tables.**

Your output should look similar to the following:

|    |cand_name|cmte_nm|
|----|----|----|
|**0**|ZUTLER, DANIEL PAUL MR|CITIZENS TO ELECT DANIEL P ZUTLER FOR PRESIDENT|
|**1**|ZUMWALT, JAMES|ZUMWALT FOR CONGRESS|
|**...**|...|...|

Consider starting from the following query skeleton, which uses the `AS` keyword to rename the `cand` and `comm` tables to `c1` and `c2`, respectively.
Which join is most appropriate?

    SELECT ...
    FROM cand AS c1
        [INNER | {LEFT |RIGHT | FULL } {OUTER}] JOIN comm AS c2
        ON ...
    ...
    ...;


```python
# %%sql
# /*
#  * Code in this scratchwork cell is __not graded.__
#  * Copy over any SQL queries you write here into the below Python cell.
#  * Do __not__ insert any new cells in between the SQL/Python cells!
#  * Doing so may break the autograder.
#  */
# -- Write below this comment. --

```

```python
query_q5a = """
...       # replace this with
...;      # your multi-line solution
"""

res_q5a = pd.read_sql(query_q5a, engine)
res_q5a
```

```python
grader.check("q5a")
```

<br/><br/>

---

### Question 5b

Suppose we modify the query from the previous part to include *all* candidates, **including those that don't have a committee.**


List the first 5 candidate names (`cand_name`) in reverse lexicographic order by `cand_name`, along with their corresponding committee names. If the candidate has no committee in the `comm` table, then `cmte_nm` should be NULL (or None in the Python representation).

Your output should look similar to the following:

|    |cand_name|cmte_nm|
|----|----|----|
|**0**|ZUTLER, DANIEL PAUL MR|CITIZENS TO ELECT DANIEL P ZUTLER FOR PRESIDENT|
|**...**|...|...|
|**4**|ZORNOW, TODD MR|None|

Hint: Start from the same query skeleton as the previous part. 
Which join is most appropriate?

```python
# %%sql
# /*
#  * Code in this scratchwork cell is __not graded.__
#  * Copy over any SQL queries you write here into the below Python cell.
#  * Do __not__ insert any new cells in between the SQL/Python cells!
#  * Doing so may break the autograder.
#  */
# -- Write below this comment. --

```

```python
query_q5b = """
...       # replace this with
...;      # your multi-line solution
"""

res_q5b = pd.read_sql(query_q5b, engine)
res_q5b
```

```python
grader.check("q5b")
```

<br/><br/>
<hr style="border: 1px solid #fdb515;" />

## Question 6: Subqueries and Grouping (OPTIONAL)

If we return to our results from Question 4, we see that many of the contributions were to the same committee:

```python
# Your SQL query result from Question 4
# reprinted for your convenience
res_q4['cmte_id'].value_counts()
```

<br/>

---

Create a new SQL query that returns the total amount that Donald Trump contributed to each committee.

Your table should have four columns: `cmte_id`, `total_amount` (total amount contributed to that committee), `num_donations` (total number of donations), and `cmte_nm` (name of the committee). Your table should be sorted in **decreasing order** of `total_amount`.

**This is a hard question!** Don't be afraid to reference the lecture slides, or the overall SQL query skeleton at the top of this lab.

Here are some other hints:

* Note that committee names are not available in `indiv_sample_nyc`, so you will have to obtain information somehow from the `comm` table (perhaps a `JOIN` would be useful).
* Remember that you can compute summary statistics after grouping by using aggregates like `COUNT(*)`, `SUM()` as output fields.
* A **subquery** may be useful to break your question down into subparts. Consider the following query skeleton, which uses the `WITH` operator to store a subquery's results in a temporary table named `donations`.

        WITH donations AS (
            SELECT ...
            FROM ...
            ... JOIN ...
                ON ...
            WHERE ...
        )
        SELECT ...
        FROM donations
        GROUP BY ...
        ORDER BY ...;

```python
# %%sql
# /*
#  * Code in this scratchwork cell is __not graded.__
#  * Copy over any SQL queries you write here into the below Python cell.
#  * Do __not__ insert any new cells in between the SQL/Python cells!
#  * Doing so may break the autograder.
#  */
# -- Write below this comment. --

```

```python
query_q6 = """
...       # replace this with
...;      # your multi-line solution
"""


res_q6 = pd.read_sql(query_q6, engine)
res_q6
```

```python
grader.check("q6")
```

# Congratulations! You finished the lab!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> res_q1.shape == (9, 1)
True
```
```
>>> "group by" not in query_q1.lower()
True
```
```
>>> sorted(list(res_q1.iloc[:,0])) == ['movie', 'short', 'tvEpisode', 'tvMiniSeries', 'tvMovie', 'tvSeries', 'tvSpecial', 'video', 'videoGame']
True
```


Tests for q2 (points: None)
```
>>> res_q2.shape == (12, 2)
True
```
```
>>> list(res_q2['total']) == [21665, 13830, 12175, 11028, 6995, 4123, 2747, 1558, 623, 410, 66, 6]
True
```
```
>>> list(res_q2['category']) == ['actor',
...  'writer',
...  'actress',
...  'producer',
...  'director',
...  'composer',
...  'cinematographer',
...  'editor',
...  'self',
...  'production_designer',
...  'archive_footage',
...  'archive_sound']
True
```


Tests for q3 (points: None)
```
>>> res_q3.shape == (10, 3)
True
```
```
>>> set(res_q3.columns) == set(['title', 'numVotes', 'averageRating'])
True
```
```
>>> np.sum(res_q3['numVotes'].astype(int)) == 19870486
True
```
```
>>> list(res_q3['title']) == ['The Shawshank Redemption',
...  'The Dark Knight',
...  'Inception',
...  'Fight Club',
...  'Pulp Fiction',
...  'Forrest Gump',
...  'Game of Thrones',
...  'The Matrix',
...  'The Lord of the Rings: The Fellowship of the Ring',
...  'The Lord of the Rings: The Return of the King']
True
```


Tests for q4 (points: None)
```
>>> res_q4.shape == (142, 3)
True
```
```
>>> sorted(res_q4['transaction_amt'].unique()) == [224, 514, 1000, 1029, 1216, 1713, 2000, 2426, 2460, 2529, 2531, 2536, 2544, 2574, 2587, 2600, 2620, 2679, 2692, 2693, 2700, 2991, 3063, 3347, 3360, 3371, 3499, 3548, 3706, 3954, 4049, 4127, 4164, 4591, 4931, 4958, 5000, 5297, 5574, 6197, 6250, 6990, 7401, 7519, 7990, 8735, 8896, 9000, 9013, 9106, 9651, 9752, 10000, 10418, 11023, 15000, 23775, 2000000, 10000000]
True
```


Tests for q5a (points: None)
```
>>> res_q5a.shape == (5, 2)
True
```
```
>>> all(res_q5a == res_q5a.sort_values('cand_name', ascending=False))
True
```
```
>>> sorted(res_q5a['cmte_nm'].unique()) == [
...     'CITIZENS TO ELECT DANIEL P ZUTLER FOR PRESIDENT',
...      'CONSTITUTIONAL COMMITTEE',
...      'JOE ZUCCOLO FOR CONGRESS',
...      'ZUKOWSKI FOR CONGRESS',
...      'ZUMWALT FOR CONGRESS']
True
```


Tests for q5b (points: None)
```
>>> res_q5b.shape == (5, 2)
True
```
```
>>> all(res_q5b == res_q5b.sort_values('cand_name', ascending=False))
True
```
```
>>> sorted(map(str, res_q5b['cmte_nm'].unique())) == [
...     'CITIZENS TO ELECT DANIEL P ZUTLER FOR PRESIDENT',
...     'JOE ZUCCOLO FOR CONGRESS',
...     'None',
...     'ZUKOWSKI FOR CONGRESS',
...     'ZUMWALT FOR CONGRESS']
True
```


Tests for q6 (points: None)
```
>>> res_q6.shape == (10, 4)
True
```
```
>>> sorted(res_q6['total_amount']) == [1000, 2000, 2600, 5000, 5000, 5200, 5400, 9000, 10000, 18633157]
True
```
```
>>> sorted(res_q6['num_donations']) == [1, 1, 1, 1, 1, 1, 1, 2, 2, 131]
True
```
```
>>> sorted(res_q6['cmte_nm']) == [
...     'DONALD J. TRUMP FOR PRESIDENT, INC.',
...     'DONOVAN FOR CONGRESS',
...     'FRIENDS OF DAVE BRAT INC.',
...     'GRASSLEY COMMITTEE INC',
...     'HELLER FOR SENATE',
...     'NEW HAMPSHIRE REPUBLICAN STATE COMMITTEE',
...     'NY REPUBLICAN FEDERAL CAMPAIGN COMMITTEE',
...     'REPUBLICAN PARTY OF IOWA',
...     'SOUTH CAROLINA REPUBLICAN PARTY',
...     'TEXANS FOR SENATOR JOHN CORNYN INC']
True
```
