```python
# Initialize Otter
import otter
grader = otter.Notebook("hw07.ipynb")
```

# Homework 7: Modeling and Analyzing COVID-19 Cases
## Probability and Estimators
## Due Date: Thursday, November 3rd, 11:59 PM PDT


**Content Warning**

This assignment includes an analysis of daily COVID-19 cases by U.S. county through 2021. If you feel uncomfortable with this topic, **please contact your GSI or the instructors.**

**Collaboration Policy**

Data science is a collaborative activity. While you may talk with others about
the homework, we ask that you **write your solutions individually**. If you do
discuss the assignments with others please **include their names** below.

**Collaborators**: *list collaborators here*



## Introduction

In this homework, we will investigate a dataset that contains information about COVID-19 cases in the United States, vaccination rates, and various other metadata that can assist in modeling various aspects of COVID-19.

Through this homework assignment, you will demonstrate your experience with:
* Bootstrap sampling
* Bias-variance tradeoff and decomposition
* Multicollinearity in features

```python
# Run this cell to set up your notebook
import numpy as np
import pandas as pd
import sklearn.linear_model as lm
from sklearn.model_selection import train_test_split

import matplotlib.pyplot as plt
import seaborn as sns
from IPython.display import Markdown

import scipy.stats

import warnings
warnings.filterwarnings("ignore")
```

<br/><br/><br/>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 1: Random Variables

Question 1 is a written problem and should be submitted as a separate PDF to the Written portal of Gradescope. All other questions in this assignment are submitted as part of this notebook.

Question 1 PDF File: https://drive.google.com/file/d/1E7iXHEY-HFxaBFt6j8upoWFMX3Mc4LX_/view?usp=share_link

<br/><br/><br/>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 2: Exploratory Data Analysis

Let's perform some initial exploratory data analysis to examine and visualize potential trends in a COVID-19 dataset.

```python
# just run this cell
covid_data = pd.read_csv('data/covid_data.csv')
covid_data.head(5)
```

The data are at county granularity; each row corresponds to COVID-19 data from one U.S. county. Here are some highlights and data sources:

* The first few columns encode county and state data; for example, check out the [FIPS](https://transition.fcc.gov/oet/info/maps/census/fips/fips.txt) numeric encoding for U.S. counties.
* The next 600 columns record daily COVID-19 cases in the county for the date range 1/22/2020 to 9/12/2021. COVID-19 case data are from CSSE at Johns Hopkins University [GitHub](https://github.com/CSSEGISandData/COVID-19/blob/master/csse_covid_19_data/csse_covid_19_time_series/time_series_covid19_confirmed_US.csv).
* The next few columns include county populations from [U.S. census data](https://www2.census.gov/programs-surveys/popest/datasets/2010-2020/counties/totals/co-est2020.csv), the latest of which is 2020.
* The last 5 columns record mask usage survey data on a 5-point scale from `NEVER` to `ALWAYS`. Data are collected in July 2020 from the New York Times [GitHub](https://github.com/nytimes/covid-19-data/blob/master/mask-use/mask-use-by-county.csv). Each column represents the proportion of population in that county who never/rarely/sometimes/frequently/always wear masks. Note, for a particular row, the numbers in those five columns sum up to $1$.

We can use `covid_data.describe()` to see various statistics about the numerical features of the provided COVID-19 data. Do any particular statistics stand out to you? Which might be useful when modeling?

**Note:** This isn't a question (i.e. it's worth no points); this is just food for thought as you start to explore the dataset.

```python
# just run this cell
covid_data.describe()
```

### Question 2a

In this homework, we will use linear regression to predict **the number of COVID-19 cases  per capita on September 12th, 2021**. Define a column `'9/12/2021_cpc'` in `covid_data` corresponding to the number of cases per capita on September 12th, 2021. 

Note that we will **always** use the `'POPESTIMATE2020'` as the population of each county.

*Hint*: The number of cases per capita should be the total number of cases in a county divided by the population of the county.


```python
...
```

```python
grader.check("q2a")
```

### Question 2b

Assign `mask_data` that has six columns from the original `covid_data` table: the five mask usage columns and the `9/12/2021_cpc` column.

**Note**: You should make a **copy** of these columns using `df.copy()` ([link](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.copy.html)).



```python
mask_data = ...
mask_data
```

```python
grader.check("q2b")
```

<!-- BEGIN QUESTION -->

### Question 2c

Our goal is to use county-wise mask usage data to predict the number of COVID-19 cases per capita on September 12th, 2021 (i.e., the column `9/12/2021_cpc`). But before modeling, let's do some EDA to explore the multicollinearality in these features, and then we will revisit this question in part 4. 

Create a visualization that shows the pairwise correlation between each combination of columns in `mask_data`. For 2-D visualizations, consider Seaborn's [heatmap](https://seaborn.pydata.org/generated/seaborn.heatmap.html). Remember to add a title to your plot.

*Hint*: You should be plotting 36 values corresponding to the pairwise correlations of the six columns in `mask_data`.


```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2d

(1) Describe the trends and takeaways visible in the visualization of pairwise correlations you plotted in Question 2c. Specifically, how does the correlation between pairs of features (i.e. mask usage) look like? How does the correlation between mask usage and cases per capita look like?

(2) If we are going to build a linear regression model (with an intercept term) using all five mask usage columns as features, what could be the problem?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/><br/>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 3: Creating a Preliminary COVID-19 Model

This question will guide you through creating a linear regression model that will predict the number of COVID-19 cases per capita given various COVID-19 safety protocols that have been implemented. Then, we will investigate the bias, variance, and observational noise of this model in the next two questions.


### Question 3a

Despite the problems we see previously, let's train a linear regression model with an intercept term, using Scikit-learn, to predict the number of COVID-19 cases per capita for September 12, 2021 using county-wise mask usage data from `mask_data`. Use `train_test_split` to evaluate your model's RMSE on a held-out test set with 33% of the COVID-19 data; call the resulting splits `X_train`, `X_test`, `y_train`, and `y_test`.

To pass the autograder, make sure to set the parameter `random_state` to 42 in your call to `train_test_split` to generate a reproducible data split ([documentation](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html)).



```python
# Create train/test sets
X = ...
y = ...
X_train, X_test, y_train, y_test = ..., ..., ..., ...

# fit the linear model and make predictions
...

# compute RMSE on train and test sets
train_rmse_cpc = ...
test_rmse_cpc = ...

...

train_rmse_cpc, test_rmse_cpc
```

```python
grader.check("q3a")
```

<!-- BEGIN QUESTION -->

### Question 3b

To visualize the model performance from part (a), let's make the following two visualizations: 

(1) the predicted values vs. observed values on the test set,

(2) the residuals plot. (Note: in multiple linear regression, the residual plot has predicted values vs. residuals) 

Some notes:
* We've used `plt.subplot` ([documentation](https://matplotlib.org/stable/gallery/subplots_axes_and_figures/subplots_demo.html)) so that you can view both visualizations side-by-side. For example, `plt.subplot(121)` sets the plottable area to the first column of a 1x2 plot grid; you can then call Matplotlib and Seaborn functions to plot that area, before the next `plt.subplot(122)` area is set.
* Remember to add a guiding line to both plot where $\hat{y} = y$, i.e., where the residual is 0.
* Remember to label your axes.


```python
plt.figure(figsize=(12,6))      # do not change this line

plt.subplot(121)                # do not change this line
# (1) predictions vs observations
...

plt.subplot(122)               # do not change this line
# (2) residual plot
...

plt.tight_layout()             # do not change this line
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3c

Describe what the plots in part (b) indicate about this linear model.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/><br/>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 4: Performing Multicollinearity Analysis

This question will guide you through performing an analysis that can reveal potential multicollinearity in our features, which is unideal. In particular, we will use bootstrap to get $95\%$ confidence intervals on the fitted parameters.

### Question 4a

Fill in the blanks below to implement the `bootstrap_sample` function, that returns $k$ randomly drawn samples from a dataset `data` of size $n$ with replacement, each of size $n$ (i.e. same size as the dataset). In other words, the returned object should be a Python list `samples` containing $k$ Pandas DataFrames, each of which have $n$ rows.

*Hint*: Take a look at the [documentation](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.sample.html) for `df.sample`!


```python
def bootstrap_sample(data, k):
    """
    Performs bootstrap sampling on data to obtain k samples of size n.
    
    Arguments:
        data - Dataset contained as a Pandas DataFrame 
        k - Number of randomly drawn samples
    
    Returns:
        samples - List containing k Pandas DataFrames of size n each
                  corresponding to each sample  
    """
    samples = []
    return ...

...
bootstrap_sample(mask_data, 1)[0]
```

```python
grader.check("q4a")
```

### Question 4b

Using the function from the previous part, let's do the following:

1. Generate 1000 bootstrapped samples from the original `mask_data` dataframe. 
2. Use Scikit-learn to fit a linear regression model (with an intercept term) to use the mask usage features to predict the `9/12/2021_cpc` response. 
3. Store each of the 1000 trained models in the Python list `models`.

Note: You *should not* create any validation or testing sets in this subpart; you should create and fit a new linear model 1000 times; you should fit your model to the entire resampled dataframe.


```python
np.random.seed(42)

datasets = ...
models = []
...

...

# These take up a lot of memory, so we should remove them!
del datasets
```

```python
grader.check("q4b")
```

### Question 4c

Fill in the blanks below in the `confidence_interval` function to generate a $95\%$ confidence interval for each of our parameters $\theta_i$, including the intercept term $\theta_0$. All of the helper code to extract coefficients from our trained models has been implemented for you already.

**Note**: 
- For a refresh on confidence intervals, refer to this link from the [Data 8 textbook](https://inferentialthinking.com/chapters/13/3/Confidence_Intervals.html). 
- To get the $i$th column from a 2D-array, you can use `2D_array[:, i]`.


```python
def extract_coefs(models, include_intercept = True):
    """
    NOTE: This function has already been implemented. You do not need to modify this!
    
    Extracts coefficients of all the linear regression models in MODELS, and returns
    it as a NumPy array with one model's coefficients as each row.
    
    Arguments:
        models - Contains k sklearn LinearRegression models, each with p + 1 coefficients
        include_intercept - Whether to include intercept in returned coefficients
    
    Returns:
        coef_array - Coefficients of all k models, each with p + 1 coefficients (if intercept
                     enabled, otherwise p). Returned object is k x (p + 1) NumPy array.
    """
    coef_array = np.zeros(shape = (len(models), len(models[0].coef_) + 1))
    for i, m in enumerate(models):
        coef_array[i, 0] = m.intercept_
        coef_array[i, 1:] = m.coef_
    if include_intercept:
        return coef_array 
    return coef_array[:, 1:]

def confidence_interval(coefs):
    """
    Calculates confidence intervals for each theta_i based on coefficients of 
    bootstrapped models. Returns output as a list of confidence intervals.
    
    Arguments:
        coefs - Output of extract_coefs, a k x (p + 1) or k x p NumPy array containing
                coefficients of bootstrapped models
    
    Returns:
        cis - Confidence intervals of each parameter theta_i in the form of a 
              list like this: [(0.5, 0.75), (0.2, 0.4), ...]
    """
    cis = []
    
    # FILL IN CODE BELOW
    for i in range(...):
        theta_i_values = ...
        theta_i_lower_ci, theta_i_upper_ci = np.percentile(...), np.percentile(...)
        cis.append((theta_i_lower_ci, theta_i_upper_ci))
    
    return cis

...

# compute confidence intervals
model_coefs = extract_coefs(models)
cis = confidence_interval(model_coefs)

# pretty print in a table
display(Markdown('#### Confidence Intervals:'))
md_list = ["|parameter|lower|upper|",
           "----|----|----|"]
md_list += [fr"|$\theta_{i}$|{lci}|{uci}|" for i, (lci, uci) in enumerate(cis)]
display(Markdown('\n'.join(md_list)))
```

```python
grader.check("q4c")
```

<!-- BEGIN QUESTION -->

### Question 4d

Interpret the confidence intervals above for each of the $\theta_i$, where $\theta_0$ is the intercept term and the remaining $\theta_i$'s are parameters corresponding to mask usage features. What does this indicate about our data and our model?

Describe a reason why this could be happening.

*Hint*: Take a look at the design matrix, heatmap, and response from Question 2!


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/><br/>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Question 5: Performing Bias-Variance Analysis

This question will guide you through performing an analysis that can estimate the bias and variance of our models, which can be helpful in modeling.

### Question 5a

We will use the same bootstrapped models contained in the Python list `models` to estimate our **model variance**. To do this, recall that the model variance on a data point is simply the variance of our predictions on that sample point. From the bias-variance decomposition in lecture, for a parametric model $\hat{Y}(x) = f_{\hat{\theta}}(x)$:

$$\text{model variance} = \mathrm{Var}(f_{\hat{\theta}}(x))$$

To investigate the variance in our test predictions, we sample a particular data point $(x_i, y_i)$. Define the **model risk** for this point as the mean square error over all possible fitted models:

$$\text{model risk} = \mathbb{E}\left[\left(y_i - f_{\hat{\theta}}(x_i)\right)^2\right]$$

Note that in contrast to lecture, you are considering a particular observation of the random response variable $Y = y_i$. Therefore model risk is an expectation over the estimate $\hat{\theta}$, which is a function of the random sample you used to fit your model.

Using the bootstrapped estimates `models`, approximate the **ratio of model variance to model risk** for the datapoint $i = 100$, i.e., $(x_{100}, y_{100})$. You can interpret this ratio as the the proportion of the expected square error on the data point "captured" by the model variance. Since you used the bootstrap to create 1000 models, you can generate 1000 predictions for the given $x_i$. Recall that `X` is the design matrix of mask features, and `y` is the `9/12/2021_cpc` response.

Assign `prop_var` to the computed, approximate ratio:

$$
\frac{\text{model variacne}}{\text{model risk}}=\frac{\mathrm{Var}\left(f_{\hat{\theta}}(x_{100})\right)}{\mathbb{E}\left[\left(y_{100} - f_{\hat{\theta}}(x_{100})\right)^2\right]}
$$


Hint: use list comprehension for creating your predictions.

```python
X_test_pt = ...
y_test_pt = ...
predictions = ...
error_squared = ...

...
preds_df = pd.DataFrame({'predictions': predictions, 'error_squared': error_squared})
prop_var = preds_df['predictions'].var() / preds_df['error_squared'].mean()


prop_var
```

```python
grader.check("q5a")
```

<!-- BEGIN QUESTION -->

### Question 5b

Comment on the ratio `prop_var`, which is the proportion of the expected square error on the data point captured by the model variance. Is the model variance the dominant term in the bias-variance decomposition? If not, what term(s) dominate the bias-variance decomposition?

**Note**: The Bias-Variance decomposition from lecture:

$$
\text{model risk} = \sigma^2 + (\text{model bias})^2 + \text{model variance}
$$

where $\sigma^2$ is the observation variance, or "irreducible error".


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 5c

Now let's calculate the average variance and average mean square error across 250 randomly sampled $(x_i, y_i)$ points. In other words, estimate the following quantities across all $x_i$ and $y_i$ in `X_sample` and `y_sample`:

$$
\frac{1}{250} \sum_{i=1}^{250} \mathrm{Var}\left(f_{\hat{\theta}}(x_i)\right)
$$

and

$$
\frac{1}{250} \sum_{i=1}^{250} \mathbb{E}\left[\left(y_i - f_{\hat{\theta}}(x_i)\right)^2\right]
$$

Hint: You can **reuse** most of your code from 5a.


```python
np.random.seed(42)

X_sample = X.sample(250)         # generate 250 x_i
y_sample = y.loc[X_sample.index] # ...and select the corresponding y_i

avg_var, avg_mse = ..., ...
...
avg_var, avg_mse
```

```python
grader.check("q5c")
```

<!-- BEGIN QUESTION -->

### Question 5d

Propose a solution to reducing the mean square error using the insights gained from the bias-variance decomposition above.

Assume that the standard bias-variance decomposition used in lecture can be applied here.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Congratulations!

Congrats! You are finished with this homework assignment.

## Detailed Submission Instructions

**There are two parts to this assignment.**
1. Question 1 is a written problem and should be submitted as a separate PDF to the Written portal of Gradescope. Please see the top of this notebook for the question writeup.

1. All other questions are submitted as part of this notebook. Please see the following cells to generate the zip file for the Coding portal of Gradescope.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q2a (points: 3)
```
>>> np.isclose(covid_data['9/12/2021_cpc'].iloc[0], 0.16541098940244012)
True
```
```
>>> np.isclose(covid_data['9/12/2021_cpc'].sum(), 397.68435739730006)
True
```


Tests for q2b (points: 1)
```
>>> mask_data.shape == (3141, 6) or mask_data.shape == (3141, 7)
True
```


Tests for q3a (points: 5)
```
>>> 0.01 <= train_rmse_cpc <= 0.05
True
```
```
>>> 0.01 <= test_rmse_cpc <= 0.05
True
```


Tests for q4a (points: 5)
```
>>> all([len(df_i) == len(mask_data) for df_i in bootstrap_sample(mask_data, 1)])
True
```
```
>>> len(bootstrap_sample(mask_data, 3)) == 3
True
```
```
>>> np.isclose(bootstrap_sample(mask_data, 1)[0]['NEVER'].mean(), 0.0799, atol = 8e-3)
True
```


Tests for q4b (points: 5)
```
>>> len(models) == 1000
True
```
```
>>> all([isinstance(model, lm.LinearRegression) for model in models])
True
```
```
>>> all([len(model.coef_) == 5 and model.intercept_ != 0 for model in models])
True
```
```
>>> all(models[0].coef_ != models[1].coef_)
True
```


Tests for q4c (points: 2)
```
>>> len(cis) == 6 and all(len(ci) == 2 for ci in cis)
True
```
```
>>> all([lci < 0 and uci > 0 for lci, uci in cis])
True
```


Tests for q5a (points: 3)
```
>>> np.isclose(prop_var, 0.001318235885796682, atol = 1e-3)
True
```


Tests for q5c (points: 3)
```
>>> avg_var <= 1e-5
True
```
```
>>> avg_mse <= 5e-3
True
```
